# Kalman filter 

This Jupyter notebook file imports (and tests) all of the code needed to do the Kalman' filter assignment. Note that the actual code is in the .py files in this directory.

This is where you'll fill in the missing code from the last assignment

Slides for this assignment: https://docs.google.com/presentation/d/1Q6w-vczvWHanGDqbuz6H1qhTOkSrX54kf1g8NgTcipQ/edit?usp=sharing

## Robot ground truth
This gets all the code in Robot_ground_truth.py. If you change any code in that file then you will need to re-execute this import.

```python
import sys; sys.path.insert(0, '.')

# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

```python
from robot_ground_truth import RobotGroundTruth, test_continuous_move_functions
```

```python
# Syntax check
robot_gt = RobotGroundTruth()
robot_gt.reset_location()
robot_gt.set_move_continuos_probabilities(0.01)
ret_amt_moved = robot_gt.move_continuous(0.0)
if -0.1 < ret_amt_moved < 0.1:
    print(f"Robot ground truth: Passed move continuous syntax check")
```

Output:
```
Robot ground truth: Passed move continuous syntax check
```

```python
# Edit the code in robot_ground_truth.py (the methods tagged with Kalman assignment)
test_continuous_move_functions()
```

Output:
```
Checking move with normal distribution probabilities
Passed continuous tests
True
```

## Robot sensors

```python
from robot_sensors import RobotSensors, test_continuous_sensor
```

```python
# Syntax check
robot_gt = RobotGroundTruth()
robot_sensor = RobotSensors()
robot_sensor.set_distance_wall_sensor_probabilities(sigma=0.01)
dist_with_noise = robot_sensor.query_distance_to_wall(robot_gt)
if 0.0 < dist_with_noise < 1.0:
    print("Dist wall sensor, passed syntax test")
```

Output:
```
Dist wall sensor, passed syntax test
```

```python
# Edit the code in robot_sensors.py (the methods tagged with Bayes assignment)
test_continuous_sensor()
```

Output:
```
Checking query wall with normal distribution probabilities
Passed continuous tests
True
```

## Kalman filter

This is the part where you actually *do* the kalman filter update. There's two parts, the sensor update and the move update. The test file runs a sensor, then a move, then a sequence of sensor, sequence of move, then a mix.

```python
from kalman_filter import KalmanFilter, test_kalman_update
```

```python
# Syntax checks
kalman_filter = KalmanFilter()
robot_ground_truth = RobotGroundTruth()
robot_sensor = RobotSensors()

# Set mu/sigmas
sensor_noise = {"mu": 0.0, "sigma": 0.1}
move_error = {"mu": 0.0, "sigma": 0.05}
robot_ground_truth.set_move_continuos_probabilities(move_error["sigma"])
robot_sensor.set_distance_wall_sensor_probabilities(sensor_noise["sigma"])

kalman_filter.update_belief_distance_sensor(robot_sensor, 0.1)
kalman_filter.update_continuous_move(robot_ground_truth, 0.0)
```

Output:
```
(0.17999999999999994, 0.13)
```

```python
test_kalman_update()
```

Output:
```
Testing Kalman
Passed
True
```

Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for kalman_filter (points: 15)
```
>>> assert(test_kalman_update(b_print=False))
```


Tests for robotGT (points: 5)
```
>>> assert(test_continuous_move_functions(b_print=False))
```


Tests for robot_sensors (points: 10)
```
>>> assert(test_continuous_sensor(b_print=False))
```
